Прежде чем проверять задания:
- **Перезапустите ядро** (**restart the kernel**) (В меню, выбрать Ядро (Kernel) $\rightarrow$ Перезапустить (Restart)
- Затем **Выполнить все ячейки**  **run all cells** (В меню, выбрать Ячейка (Cell) $\rightarrow$ Запустить все (Run All).

Убедитесь, что заполнены все ячейки с комментарием "НАЧАЛО ВАШЕГО РЕШЕНИЯ".

После ячеек с заданием следуют ячейки с проверкой с помощью assert.

Если в коде есть ошибки, assert выведет уведомление об ошибке.

Если в коде нет ошибок, assert отработает без вывода дополнительной информации.

---

## Импорт библиотек

```python
import numpy as np
```

# Задания на векторы

## Задание 1 (1 балл)

Напишите код функции для вычисления L2-нормы одномерный np.array вектора с помощью numpy:

<!-- grade_id: l2_norm, points: 0, grade: false, solution: true -->
```python
def l2_norm(vector):
    
    """
    Аргументы:
    vector - одномерный np.array вектор

    Возвращает:
    float
    """
    
    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    norm = np.linalg.norm(vector)
    return norm
    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: l2_norm_test, points: 1, grade: true, solution: false -->
```python
# Тест 1: Вектор (1, 0, 0) имеет L2-норму 1.
vector1 = np.array([1, 0, 0])
assert l2_norm(vector1) == 1.0

# Тест 2: Вектор (0, 1, 0) имеет L2-норму 1.
vector2 = np.array([0, 1, 0])
assert l2_norm(vector2) == 1.0

# Тест 3: Вектор (1, 2, 3, 4, 5, 6, 7, 8, 9) имеет L2-норму sqrt(285).
vector3 = np.array([1, 2, 3, 4, 5, 6, 7, 8, 9])
assert l2_norm(vector3) == np.sqrt(285)

# Тест 3: Вектор (0, 0, 1) имеет L2-норму 1.
vector3 = np.array([0, 0, 1])
assert l2_norm(vector3) == 1.0

# Тест 4: Вектор (3, 4) имеет L2-норму 5.
vector4 = np.array([3, 4])
assert l2_norm(vector4) == 5.0

# Тест 5: Вектор (-3, -4) имеет L2-норму 5.
vector5 = np.array([-3, -4])
assert l2_norm(vector5) == 5.0

# Тест 6: Вектор (1, 1, 1) имеет L2-норму sqrt(3).
vector6 = np.array([1, 1, 1])
assert l2_norm(vector6) == np.sqrt(3)

# Тест 7: Вектор (0, 0, 0) имеет L2-норму 0.
vector7 = np.array([0, 0, 0])
assert l2_norm(vector7) == 0.0

# Тест 8: Вектор (1, 2, 3, 4, 5) имеет L2-норму sqrt(55).
vector8 = np.array([1, 2, 3, 4, 5])
assert l2_norm(vector8) == np.sqrt(55)

# Тест 9: Вектор (1, 2, 3, 4, 5, 6) имеет L2-норму sqrt(91).
vector9 = np.array([1, 2, 3, 4, 5, 6])
assert l2_norm(vector9) == np.sqrt(91)

# Тест 10: Вектор (1, 2, 3, 4, 5, 6, 7, 8, 9) имеет L2-норму sqrt(285).
vector10 = np.array([1, 2, 3, 4, 5, 6, 7, 8, 9])
assert l2_norm(vector10) == np.sqrt(285)
```

---
## Задание 2 (1 балл)

Напишите код функции, которая на вход принимает два одномерных np.array массива чисел и производит их сложение с помощью numpy. Если длины векторов не равны, функция должна вернуть None

<!-- grade_id: add_vectors, points: 0, grade: false, solution: true -->
```python
def add_vectors(arr1, arr2):
    """
    Складывает два одномерных массива.

    Аргументы:
    arr1, arr2 -- массивы, представленные в виде np.array

    Возвращает:
    Массив, который является результатом сложения
    """
    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    if len(arr1) == len(arr2):
        return np.add(arr1, arr2)
    else:
        return None
    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: add_vectors_test, points: 1, grade: true, solution: false -->
```python
arr1 = np.array([1, 2, 3])
arr2 = np.array([4, 5, 6])
assert np.array_equal(add_vectors(arr1, arr2), np.array([5, 7, 9]))

arr1 = np.array([1.5, 2.0, 3.5])
arr2 = np.array([4.2, 5.7, 6.1])
assert np.allclose(add_vectors(arr1, arr2), np.array([5.7, 7.7, 9.6]))

arr1 = np.array([])
arr2 = np.array([])
assert np.array_equal(add_vectors(arr1, arr2), np.array([]))

arr1 = np.array([-1, 2, -3])
arr2 = np.array([4, -5, 6])
assert np.array_equal(add_vectors(arr1, arr2), np.array([3, -3, 3]))

arr1 = np.array([1, 2, 3])
arr2 = np.array([4])
assert np.array_equal(add_vectors(arr1, arr2), None)

arr1 = np.array([1, 2, 3])
arr2 = np.array([4, 5])
assert np.array_equal(add_vectors(arr1, arr2), None)

arr1 = np.array([1, 2, 3, 4, 5])
arr2 = np.array([2, 4, 6, 8, 10])
assert np.array_equal(add_vectors(arr1, arr2), np.array([3, 6, 9, 12, 15]))

arr1 = np.array([1, 2, 3])
arr2 = np.array([0, 0, 0])
assert np.array_equal(add_vectors(arr1, arr2), arr1)

arr1 = np.array([1, 2, 3])
arr2 = np.array([-1, -2, -3])
assert np.array_equal(add_vectors(arr1, arr2), np.array([0, 0, 0]))

arr1 = np.array([1, 2, 3])
arr2 = np.array([1, 2, 3])
assert np.array_equal(add_vectors(arr1, arr2), np.array([2, 4, 6]))
```

---
## Задание 3 (1 балл)

Напишите код функции, которая на вход принимает два одномерных np.array массива чисел и вычисляет их скалярное произведение с помощью numpy. Если длины векторов не равны или векторы пустые, необходимо вернуть None

<!-- grade_id: scalar_product, points: 0, grade: false, solution: true -->
```python
def scalar_product(arr1, arr2):
    """
    Вычисляет скалярное произведение двух одномерных массивов.

    Аргументы:
    arr1, arr2 -- массивы, представленные в виде np.array

    Возвращает:
    Скаляр, который является результатом вычисления
    """

    
    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    if arr1.size != arr2.size or (arr1.size == 0 or arr2.size == 0) :
        return None
    else:
        return np.dot(arr1, arr2)
    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: scalar_product_test, points: 1, grade: true, solution: false -->
```python
a = np.array([1, 2, 3])
b = np.array([4, 5, 6])
assert scalar_product(a, b) == 32

a = np.array([1, 0, 0])
b = np.array([0, 1, 0])
assert scalar_product(a, b) == 0

a = np.array([1, 2, 3, 4, 5])
b = np.array([5, 4, 3, 2, 1])
assert scalar_product(a, b) == 35

a = np.array([1, 2])
b = np.array([4, 5, 6])
assert scalar_product(a, b) is None

a = np.array([-1, 2, -3])
b = np.array([4, -5, 6])
assert scalar_product(a, b) == -32

a = np.array([0, 0, 0])
b = np.array([4, 5, 6])
assert scalar_product(a, b) == 0

a = np.array([1, 2, 3])
b = np.array([4, 5])
assert scalar_product(a, b) is None

a = np.array([])
b = np.array([4, 5, 6])
assert scalar_product(a, b) is None

a = np.array([])
b = np.array([])
assert scalar_product(a, b) is None

a = np.array([1])
b = np.array([4])
assert scalar_product(a, b) == 4
```

---
## Задание 4 (1 балл)

Напишите код функции, которая на вход принимает два одномерных np.array массива чисел и вычисляет косинусное расстояние между ними с помощью numpy. Если длины векторов не равны, или один из векторов пустой, необходимо вернуть None. В случае, если один из векторов нулевой, также нужно вернуть None

<!-- grade_id: cosine_distance, points: 0, grade: false, solution: true -->
```python
def cosine_distance(arr1, arr2):
    """
    Вычисляет косинусное расстояние между двумя одномерными массивами.

    Аргументы:
    arr1, arr2 -- массивы, представленные в виде np.array

    Возвращает:
    float, который является результатом вычисления
    """

    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    if arr1.size != arr2.size or (arr1.size == 0 or arr2.size == 0) or (not np.any(arr1)) or (not np.any(arr2)):
        return None
    else:
        cosine_distance = 1 - np.dot(arr1, arr2) / (np.linalg.norm(arr1) * np.linalg.norm(arr2))
        return cosine_distance
    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: cosine_distance_test, points: 1, grade: true, solution: false -->
```python
assert np.isclose(cosine_distance(np.array([1, 2, 3]), np.array([4, 5, 6])), 0.025368153802923787)
assert cosine_distance(np.array([1, 2, 3]), np.array([4, 5])) is None
assert cosine_distance(np.array([1, 2, 3]), np.array([])) is None
assert np.isclose(cosine_distance(np.array([0, 1, 0, 1]), np.array([1, 0, 1, 0])), 1.0)
assert np.isclose(cosine_distance(np.array([1, 0, 1]), np.array([-1, 0, -1])), 2.0)
assert np.isclose(cosine_distance(np.array([3, 4]), np.array([6, 8])), 0.0)
assert cosine_distance(np.array([]), np.array([1, 2, 3])) is None
assert cosine_distance(np.array([]), np.array([])) is None
assert cosine_distance(np.array([1, 2, 3]), np.array([0, 0, 0])) is None
assert cosine_distance(np.array([0, 0, 0]), np.array([1, 2, 3])) is None
```

# Задания на матрицы

---
## Задание 5 (1 балл)

Напишите функцию, которая на вход принимает два np.array массива чисел и производит их поэлементное сложение с помощью numpy. Если размеры матриц не совпадают или массивы пустые (не содержат элементов), вернуть None

<!-- grade_id: add_arrays, points: 0, grade: false, solution: true -->
```python
def is_empty(arr):
    """
    Вспомогательная функция
    Проверяет, является ли многомерный массив пустым.

    Аргументы:
    arr -- массив, представленный в виде np.array

    Возвращает:
    bool -- True, если массив пустой, False в противном случае
    """
    return arr.size == 0



def add_arrays(arr1, arr2):
    """
    Производит поэлементное сложение двух массивов.

    Аргументы:
    arr1, arr2 -- массивы, представленные в виде np.array

    Возвращает:
    np.array, являющийся результатом сложения
    """
    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    if arr1.size != arr2.size or is_empty(arr1) or is_empty(arr2):
        return None
    else:
        return np.add(arr1, arr2)
    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: add_arrays_test, points: 1, grade: true, solution: false -->
```python
arr1 = np.array([1, 2, 3])
arr2 = np.array([4, 5, 6])
assert np.array_equal(add_arrays(arr1, arr2), np.array([5, 7, 9]))

arr1 = np.array([1, 2, 3])
arr2 = np.array([])
assert add_arrays(arr1, arr2) is None

arr1 = np.array([1.1, 2.2, 3.3])
arr2 = np.array([4.4, 5.5, 6.6])
assert np.allclose(add_arrays(arr1, arr2), np.array([5.5, 7.7, 9.9]))

arr1 = np.array([1, 2, 3])
arr2 = np.array([4, 5, 6])
assert np.array_equal(add_arrays(arr1, arr2), add_arrays(arr2, arr1))

arr1 = np.array([1, 2, 3])
arr2 = np.array([4, 5, 6])
arr3 = np.array([7, 8, 9])
assert np.array_equal(add_arrays(add_arrays(arr1, arr2), arr3), add_arrays(arr1, add_arrays(arr2, arr3)))

arr1 = np.array([1, 2, 3])
arr2 = np.array([4, 5, 6])
assert np.array_equal(add_arrays(arr1, arr2), add_arrays(arr1, arr2))

arr1 = np.array([])
arr2 = np.array([1, 2, 3])
assert add_arrays(arr1, arr2) is None

arr1 = np.array([])
arr2 = np.array([[], []])
assert add_arrays(arr1, arr2) is None

arr1 = np.random.rand(1000)
arr2 = np.random.rand(1000)
assert np.allclose(add_arrays(arr1, arr2), arr1 + arr2)
```

---
## Задание 6 (1 балл)

Напишите функцию, которая на вход принимает два np.array массива чисел и производит их умножение с помощью numpy. Должны соблюдаться правила произведения матриц. Если правила не соблюдаются, функция должна вернуть None

<!-- grade_id: matrix_multiply, points: 0, grade: false, solution: true -->
```python
import numpy as np

def matrix_multiply(a, b):
    """
    Умножает две матрицы a и b с помощью numpy.

    Аргументы:
    a, b -- две матрицы, представленные в виде np.array

    Возвращает:
    Матрицу, полученную в результате умножения матриц a и b
    """
    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    if (a.ndim > 1 and a.shape[1] == b.shape[0]) or (b.ndim == a.size):
            return np.dot(a, b)
    else:
            return None  
    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: matrix_multiply_test, points: 1, grade: true, solution: false -->
```python
a1 = np.array([[1, 2], [3, 4]])
b1 = np.array([[5, 6], [7, 8]])
c1 = np.array([[19, 22], [43, 50]])
assert np.array_equal(matrix_multiply(a1, b1), c1)

a2 = np.array([[1, 2, 3], [4, 5, 6]])
b2 = np.array([[7, 8], [9, 10], [11, 12]])
c2 = np.array([[58, 64], [139, 154]])
assert np.array_equal(matrix_multiply(a2, b2), c2)

a3 = np.array([[1, 2, 3], [4, 5, 6]])
b3 = np.array([[7, 8], [9, 10]])
assert matrix_multiply(a3, b3) is None

a4 = np.array([1, 2, 3])
b4 = np.array([4, 5, 6])
assert matrix_multiply(a4, b4) is None

a5 = np.array([[1, 2], [3, 4]])
b5 = np.array([[5], [6]])
c5 = np.array([[17], [39]])
assert np.array_equal(matrix_multiply(a5, b5), c5)

a6 = np.array([[1, 2, 3]])
b6 = np.array([[4], [5], [6]])
c6 = np.array([[32]])
assert np.array_equal(matrix_multiply(a6, b6), c6)

a7 = np.array([[]])
b7 = np.array([[1]])
assert matrix_multiply(a7, b7) is None

a8 = np.array([[]])
b8 = np.array([[]])
assert matrix_multiply(a8, b8) is None
```

---
## Задание 7 (1 балл)

Напишите функцию, которая на вход принимает np.array массив чисел и вычисляет определитель с помощью numpy. Если матрица не квадратная, функция должна вернуть None

<!-- grade_id: determinant, points: 0, grade: false, solution: true -->
```python
def determinant(arr):
    """
    Вычисляет определитель матрицы arr с помощью numpy.

    Аргументы:
    arr -- матрица, представленная в виде np.array

    Возвращает:
    Определитель матрицы, либо None
    """    
    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    if (arr.ndim > 1 and arr.shape[0] == arr.shape[1]) or (arr.ndim == 1 and arr.size == 1):
        return np.linalg.det(arr)
    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: determinant_test, points: 1, grade: true, solution: false -->
```python
arr1 = np.array([[1, 2], [3, 4]])
assert np.isclose(determinant(arr1), -2)

arr2 = np.array([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12]])
assert determinant(arr2) == None

arr3 = np.array([[1]])
assert np.isclose(determinant(arr3), 1)

arr2 = np.array([[1, 0, 0], [0, 1, 0], [0, 0, 1]])
assert np.isclose(determinant(arr2), 1)

arr3 = np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]])
assert np.isclose(determinant(arr3), 0)

arr5 = np.array([[1, 0], [0, 1], [0, 0]])
assert determinant(arr5) == None

arr6 = np.array([])
assert determinant(arr6) == None

arr8 = np.array([[1, 2], [3, 4], [5, 6]])
assert determinant(arr8) == None

arr9 = np.array([[1, 2, 3], [4, 5, 6]])
assert determinant(arr9) == None

arr10 = np.array([[1, 2], [3, 4], [5, 6]])
assert determinant(arr10) == None
```

# Задание на производные

---
## Задание 8 (1 балл)

Напишите функция для вычисления производной функции $\frac{7x^3}{3} - \frac{5x^2}{2}$ в точке x. Функция должна возврвращать переменную df_dx

<!-- grade_id: calculate_dderivative, points: 0, grade: false, solution: true -->
```python
df_dx = 0
def calculate_dderivative(x: float):
    """
    Вычисляет производную \frac{7x^3}{3} - \frac{5x^2}{2} в точке x

    Аргументы:
    x -- float
    
    Возвращает:
    df_dx - производная \frac{7x^3}{3} - \frac{5x^2}{2} в точке x
    """
    # НАЧАЛО ВАШЕГО РЕШЕНИЯ
    def f(x):
        return (7 / 3 * x**3 - 5 / 2 * x**2)
    h = 0.001
    df_dx = (f(x+h) - f(x))/h
    # КОНЕЦ ВАШЕГО РЕШЕНИЯ
    return df_dx
```

<!-- grade_id: calculate_dderivative_test, points: 1, grade: true, solution: false -->
```python
assert np.isclose(calculate_dderivative(0), 0, atol=0.1)
assert np.isclose(calculate_dderivative(1), 2, atol=0.1)
assert np.isclose(calculate_dderivative(2), 18, atol=0.1)
assert np.isclose(calculate_dderivative(3), 48, atol=0.1)
assert np.isclose(calculate_dderivative(5), 150, atol=0.1)
```